---
id: 6577c17a1c096764b6139e82
title: priority-party-planner-task-count
challengeType: 0
dashedName: priority-party-planner-task-count
---

# --description--

In the Priority Party Planner challenge, you're tasked with creating a function 'planPartyTasks' that organizes party tasks based on their priority and duration. In this step, we'll start by ensuring your function can accept an array of the task objects and return the correct task count.

Each task object in the array has properties: 'id', 'priority', and 'duration'.

Example task object:
{id: 'task1', priority: 1, duration: 30}

# --hints--

'planPartyTasks([{id: 'task1', priority: 1, duration: 30}])' should return a number


```js
assert(typeof planPartyTasks([{ id: 'task1', priority: 1, duration: 30 }]) === 'number');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
